Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

---

```python
from sklearn.datasets import load_diabetes
import numpy as np
from sklearn.kernel_ridge import KernelRidge

def rel_error(x, y):
  """ returns relative error """
  return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))

data = load_diabetes()
X_train, y_train = data.data, data.target
```

<!-- grade_id: cell-d36cbb79389907e8, points: 0, grade: false, solution: true -->
```python
def gaussian_kernel(x1, x2, sigma=1.0):
    # YOUR CODE HERE
    return np.exp(-np.linalg.norm(x1-x2)**2/(2*sigma**2))

class KernelRidgeRegression():
    def __init__(self, kernel, lambd=1.0):
        self.lambd = lambd # regularization coefficient
        self.alpha = None # dual variable
        self.kernel = kernel # kernel function
    
    def fit(self, X, y):
        # YOUR CODE HERE
        N = X.shape[0]
        I = np.eye(N)
        
        #Kernel function
        k = np.zeros((N,N))
        for i in range(N):
            for j in range(N):
                k[i][j] = self.kernel(X[i,:],X[j,:])
        
        temp = np.linalg.inv(k+self.lambd*I)
        self.alpha = y@temp     
    
    def predict(self, X):
        # YOUR CODE HERE
        N = X.shape[0]
        y_pred = np.zeros(N)
        for i in range(N):
            k = np.zeros(N)
            for j in range(N):
                k[j] = self.kernel(X[i,:],X[j,:])
            y_pred[i] = self.alpha@k
        return y_pred    
```

<!-- grade_id: cell-ec40f17b949370cd, points: 5, grade: true, solution: false -->
```python
kr = KernelRidgeRegression(kernel=gaussian_kernel)
kr.fit(X_train,y_train)
y_pred = kr.predict(X_train)

sklearn_kr = KernelRidge(kernel=gaussian_kernel)
sklearn_kr.fit(X_train,y_train)
sklearn_y_pred = kr.predict(X_train)

error = rel_error(kr.alpha, sklearn_kr.dual_coef_)
print("error: ", error)
assert error <= 1e-11

error_pred = rel_error(y_pred, sklearn_y_pred)
print("prediction error: ", error_pred)
assert error <= 1e-11
```

Output:
```
error:  7.175796734076663e-13
prediction error:  0.0
```

```python
kr = KernelRidgeRegression(kernel=gaussian_kernel)
kr.fit(X_train,y_train)
y_pred = kr.predict(X_train)

sklearn_kr = KernelRidge(kernel=gaussian_kernel)
sklearn_kr.fit(X_train,y_train)
sklearn_y_pred = sklearn_kr.predict(X_train)

error = rel_error(kr.alpha, sklearn_kr.dual_coef_)
print("error: ", error)
assert error <= 1e-11

error_pred = rel_error(y_pred, sklearn_y_pred)
print("prediction error: ", error_pred)
assert error <= 1e-11
```

Output:
```
error:  7.175796734076663e-13
prediction error:  8.821735134558615e-15
```

```python

```